```python
# Initialize Otter
import otter
grader = otter.Notebook("HW1_Fall2022.ipynb")
```

# DATA 601

## HW1


**Learning Objectives**

- Explore built-in data types in Python.
- Review fundamental programming and problem solving concepts with Python.
- Implement functions based on mathematical concepts and definitions.
- Gain experience working with the Jupyter notebook environment.

_This is an individual homework assignment._ 

Please complete this homework assignment within the Jupypter notebook environment.  

#### Submission 

In order to ensure that everything goes smoothly, and is easy to grade please follow these instructions:

- Please provide your solutions where asked; please do not alter any other parts of this notebook.


#### <font color='red'> Note: Passing the provided test-cases doesn't guarantee points. We check your code with other test-cases for grading.</font> 

```python
# Check that we are using a recent version of Jupyter.
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
import otter
grader = otter.Notebook("HW1_Fall2022.ipynb")
```

## Part A

This part focuses on scalar types. You should be able to complete the following questions without using any collection types.

### Question 1

The central bionomial coefficients form a vertical row down the middle of Pascal's triangle. They can be calculated via the following formula:

$$
{2n\choose n} = \frac{(2n)!}{(n!)^2} = \prod_{k=1}^{n} \frac{n+k}{k}  
$$

Write a function that calculates the nth central bionomial coefficient. You may use in-built functions to calculate the factorials.

(3 points)

```python
import math 
def cent_binomial_coef(n):
    result = math.comb(2*n,n)
    return result
    '''Returns the n-th binomial coefficient for an input
    integer n.
    '''
    ...
```

```python
grader.check("q1_a")
```

Output:
```
q1_a results: All test cases passed!
```

```python
cent_binomial_coef(0)
```

Output:
```
1
```

### Question 2

Write a function that determines whether an integer is prime. You may not use built-in functions to determine this directly.

(8 points)

```python
def is_prime(n):
    diviser = 2
    while diviser < n:
        if (n % diviser) == 0:
            return False
        diviser += 1
    return True
    '''For a positive integer n, returns True if n is a 
    prime number, False otherwise.
    '''
    ...
```

```python
grader.check("q2_a")
```

Output:
```
q2_a results: All test cases passed!
```

##### Question 3

(5 points)

Many things in nature can be modeled by the exponential growth function: 
$$
P(t) = \alpha \beta^{\frac{t}{\tau}},
$$
where $\alpha$ is the initial population, $\beta$ is the growth rate factor, t is the time passed, and $\tau$ is the time constant. 

Imagine that you are an entomologist on the plant Zergo. There is a small insect of interest to you on this planet which has an exponential growth rate. From experimentation you determine that: 
$\beta= \sqrt\frac{5}{4}$,  and $\tau= 2^{\frac{1}{3}}$.

Starting with a population $\alpha$ of 10, use this formula to compute $P(t)$ in the function `population(t)` below. 

<!--
BEGIN QUESTION
name: A3
manual: true
points: 2
-->

```python
import math
def population(t):
    initial = 10
    pop = math.floor(initial*(math.sqrt(5/4)**(t/(2**(1/3)))))
    return pop
    
    '''Computes the population at time t, where t is a positive float, 
    and returns the population as an integer (floor value). 
    It is assumed that time is in minutes and that all parameters
    are correct for this assumption.
    '''
    ...
```

```python
grader.check("q3_a")
```

Output:
```
q3_a results: All test cases passed!
```

## Part B

Questions in this part rely on strings, lists, and tuples.

### Question 1

Consider the lists $a=[2,1,2,5,3,1]$ and $b=[2,1,7,1]$. The intersection of these lists, $a \cap b$, is $[1,1,2]$. Write a function that returns the intersection of two arbitrary lists. You may not use built-in set operations, and should pay attention to duplicates.

(5 points)


```python
def intersection_list(a,b):
    intersection = []
    for i,u in enumerate(a):
        for x,y in enumerate(b):
            try:
                if u == y:
                    intersection.append(y)
                    b.remove(x)
                    a.remove(i)
                    pass
            except:
                pass
    return intersection
    '''
    Returns the intersection of two lists a and b. 
    
    '''
    ...
```

```python
grader.check("q1_b")
```

Output:
```
q1_b results: All test cases passed!
```

### Question 2
(8 points)

The scalar triple product of three vectors is defined as:

$$
a.(b \times c)
$$

Write a function that calculates the scalar triple product of three vectors, a, b, and c.

```python
def striple_prod(a,b,c):
    x = b[1]*c[2] - b[2]*c[1]
    y = b[2]*c[0] - b[0]*c[2]
    z = b[0]*c[1] - b[1]*c[0]
    result = (a[0]*x)+(a[1]*y)+(a[2]*z)   
    return result
    '''Returns the scalar triple product of three vectors A, B, and C. 
    '''
    ...
```

```python
grader.check("q2_b")
```

Output:
```
q2_b results: All test cases passed!
```

### Question 3
(10 points)

Write a function that, given a piece of text, calculates the most frequently occuring letter.


```python
import collections
def freq_letters(text):
    text = text.lower()
    alphabet = ['a','b','c','d','e','f','g','h','i','j','k','l','m','n','o','p','q','r','s','t','u','v', 'w','x','y','z']
    d = collections.defaultdict(int)
    #initialize dict
    for i in text:
        if i in alphabet:
            d[i] += 1
    return sorted(d.items(), key=lambda x: x[1], reverse=True)[0][0]
    '''
    For textual input, returns the most frequently occuring letter 
    (your function should be case-insensitive)
    '''
    ...
```

```python
grader.check("q3_b_1")
```

Output:
```
q3_b_1 results: All test cases passed!
q3_b_1 - 1 message: Good job!
```

Write a second function that returns the number of vowels and consonants in the text.

```python
def vowels(text):
    num_vowels = 0
    num_consonants = 0
    vowels =['a','e','i','o','u','y']
    consonants = ['b','c','d','f','g','j','k','l','m','n','p','q','s','t','v','x','z']
    for i in text:
        for x in vowels:
            if x == i.lower():
                num_vowels += 1
        for y in consonants:
                  if y == i.lower():
                      num_consonants += 1
    return (num_vowels, num_consonants)
    '''
    For textual input, returns the number of vowels and consonants.
    Return a tuple with number of vowels as the first element and consonants
    as the second element, i.e (num_vowels, num_consonant)
    '''
    ...
```

```python
grader.check("q3_b_2")
```

Output:
```
q3_b_2 results: All test cases passed!
```

## Part C

The question in this part of the assignment is intended to:

  1) Illuminate how what you already know is useful in this context.
  
  2) Provide an exercise in writing about coding concepts and ideas and using the relevant language and terminology. 
  
  3) Demonstrate your understanding of the assignment material in a different format.


<!-- BEGIN QUESTION -->

### Question 1 
(3 points) 

In 100 to 250 words identify a concept you have found difficult or confusing from this assignment. Reflect on how your previous learning or experience helped you to understand this concept. Provide your reflection using markdown in the cell below. 

_I had the most trouble with the intersection in Question 1 of Part B. The main reason was because I was creating nested loops thhat were becoming difficult to keep track of. The previous experience that helped me was tracing my code step-by-step as well as asking myself what each condition/loop was accomplishing. This allowed me to break the function down and remember what should be in each loop._

<!-- END QUESTION -->

## Submission

Make sure you have run all cells in your notebook and that your code works as expected. Save your notebook and upload it to <a href=https://www.gradescope.ca >Gradescope</a>
**Please save (ctrl+s) before uploading!**

```python

```

Tests for q1_a (points: 3)
```
>>> cent_binomial_coef(0) == 1
True
```
failure_message: Don't forget (0)! is 1.
```
>>> cent_binomial_coef(1) == 2
True
```


Tests for q1_b (points: 5)
```
>>> def testcase1(intersection_list):
...     temp = intersection_list([2,1,2,5,3,1], [2,1,7,1])
...     length = len(temp)
...     ones = temp.count(1)
...     twos = temp.count(2)
...     assert (length, ones, twos) == (3, 2, 1)
>>> testcase1(intersection_list)
```


Tests for q2_a (points: 8)
```
>>> is_prime(2) == True
True
```
```
>>> is_prime(3) == True
True
```
```
>>> is_prime(4) == False
True
```


Tests for q2_b (points: 8)
```
>>> striple_prod([5, 1, 2], [-1]*3, [2, 1, 0]) == 5
True
```


Tests for q3_a (points: 5)
```
>>> population(5.7) == 16
True
```
```
>>> population(2.9) == 12
True
```


Tests for q3_b_1 (points: 5)
```
>>> freq_letters('a!!!!!').lower() == 'a'
True
```
failure_message: You should only count letters a-z and not numbers, punctuation, spaces, etc.
success_message: Good job!


Tests for q3_b_2 (points: 5)
```
>>> vowels('I like banana!!') == (6, 5)
True
```
